---
title: Template Chapter 1
description: This is a template chapter.
---

## Ex 1.1

```yaml
type: NormalExercise
lang: r
xp: 100
skills: 1
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```

`@sample_code`

```{r}

```

`@solution`

```{r}

```

`@sct`

```{r}

```
